# MLflow Exercise Background

<img src="images/overview-mlflow-focus.jpg" width=800>

As shown in the figure above, the MLOps platform provides an MLflow service that stores ML model training information to a PostgreSQL database and training artifacts to a MinIO storage service. 

MLflow provides a [Python client](https://mlflow.org/docs/latest/python_api/index.html) for communicating with an MLflow service. For example, we can use the MLflow Python client to start an *MLflow run*, which is an execution of an ML training script:
```python
with mlflow.start_run():
    model = ElasticNet(alpha=..., l1_ratio=...)
    model.fit(train_x, train_y)
```
*MLflow runs* are organized into *MLflow experiments*. An MLflow experiment can be seen as a logical unit of one or more MLflow runs. For example, there can be an MLflow experiment for training an ElasticNet model, and there can be multiple MLflow runs under this experiment for exploring different hyperparameters and/or training datasets.

When starting  an MLflow run, we can record the relevant information, such as the configured hyperparameters and custom evaluation metrics. After the run is completed, we can also upload the produced model artifact to MLflow:
```python
with mlflow.start_run():
    model = ElasticNet(alpha=..., l1_ratio=...)
    model.fit(train_x, train_y)
    mlflow.log_param("alpha", ...)
    mlflow.log_param("l1_ratio", ...)
    mlflow.log_metric("rmse", ...)
    mlflow.sklearn.log_model(model, ...)
```
Below is a complete example. 

More reading material: [MLflow docs](https://mlflow.org/docs/latest/index.html)

# MLflow example
In this example, we'll use sklearn to train a simple ElasticNet model that predicts red wine quality given some chemical attributes. The information of dataset used in this example can be found [here](https://archive.ics.uci.edu/dataset/186/wine+quality). 

# Take a look at the dataset

```python
from deepchecks.tabular.datasets.regression import wine_quality
from deepchecks.tabular import Dataset
from sklearn.model_selection import train_test_split
from sklearn.ensemble import RandomForestRegressor
import pandas as pd

# Load model

#train, test = wine_quality.load_data(as_train_test=True) # set to false to get one big chunk of data
train_data = pd.read_csv('data/1_data_train.csv')
test_data = pd.read_csv('data/1_data_test.csv')
target = 'quality'

target_train = train_data[target]
train_data = train_data.drop(target, axis=1)
X_train = train_data
y_train = target_train

test_data = test_data.drop(target, axis=1)
y_test = test_data[target]
X_test = test_data

train_ds = Dataset(X_train, label=y_train, cat_features=[])
test_ds  = Dataset(X_test,  label=y_test,  cat_features=[])

#train_data.nlargest(10, 'pH') # Get top 10 rows with largest pH (or any other column)
X_train.nlargest(10, 'alcohol') # Get top 10 rows with largest pH (or any other column)
```

### Create an MLflow run
The following code snippet exemplifies how to use the MLflow Python client to record training parameters and evaluation metrics as well as upload the trained model artifact to the MLflow service.

```python
import os
import logging

import mlflow
import mlflow.sklearn
import numpy as np
import pandas as pd
from sklearn.metrics import mean_squared_error
from sklearn.model_selection import train_test_split
from sklearn.linear_model import ElasticNet

# Set an environmental variable named "MLFLOW_S3_ENDPOINT_URL" so that MLflow client knows where to save artifacts.
# The MinIO storage service can be accessed via http://mlflow-minio.local
os.environ["MLFLOW_S3_ENDPOINT_URL"] = "http://mlflow-minio.local"

# Configure the credentials needed for accessing the MinIO storage service.
# "AWS_ACCESS_KEY_ID" has been configured in a ComfigMap and "AWS_SECRET_ACCESS_KEY" in a Secret in your K8s cluster when you set up the MLOps platform
os.environ["AWS_ACCESS_KEY_ID"] = "minioadmin"
os.environ["AWS_SECRET_ACCESS_KEY"] = "minioadmin"

logging.basicConfig(level=logging.INFO)
logger = logging.getLogger(__name__)

MLFLOW_TRACKING_URI = "http://mlflow-server.local" # This is the URL of the MLflow service
#MLFLOW_EXPERIMENT_NAME = "YOUR EXPERIMENT NAME HERE"
#YOUR_MODEL_NAME="REPLACE_WITH_YOUR_MODEL_NAME"


def eval_metrics(actual, pred):
    rmse = np.sqrt(mean_squared_error(actual, pred))
    return rmse


def run_mlflow_example():
    np.random.seed(40)
    # Just use hard-coded hyperparameters
    alpha = 0.5
    l1_ratio = 0.5

    logger.info(f"Using MLflow tracking URI: {MLFLOW_TRACKING_URI}")

    # Configure the MLflow client to connect to the MLflow service
    mlflow.set_tracking_uri(MLFLOW_TRACKING_URI)

    logger.info(f"Using MLflow experiment: {MLFLOW_EXPERIMENT_NAME}")
    mlflow.set_experiment(MLFLOW_EXPERIMENT_NAME)

    with mlflow.start_run() as run:
        print("MLflow run_id:", run.info.run_id) # Each MLflow Run has a unique identifier called run_id

        lr = ElasticNet(alpha=alpha, l1_ratio=l1_ratio, random_state=42)

        logger.info("Fitting model...")

        lr.fit(X_train, y_train)

        logger.info("Finished fitting")

        predicted_qualities = lr.predict(X_test)

        rmse = eval_metrics(y_test, predicted_qualities)

        logger.info("Elasticnet model (alpha=%f, l1_ratio=%f):" %
                    (alpha, l1_ratio))
        logger.info("  RMSE: %s" % rmse)


        logger.info("Logging parameters to MLflow")
        mlflow.log_param("alpha", alpha)
        mlflow.log_param("l1_ratio", l1_ratio)
        mlflow.log_metric("rmse", rmse)

        logger.info("Logging trained model")
        artifact_name = "model"
        mlflow.sklearn.log_model(
            lr, artifact_name, registered_model_name=YOUR_MODEL_NAME)
        print("The S3 URI of the logged model:", mlflow.get_artifact_uri(artifact_path=artifact_name))
```

```python
run_mlflow_example()
```

Output:
```
INFO:__main__:Using MLflow tracking URI: http://mlflow-server.local
INFO:__main__:Using MLflow experiment: mlflow-minio-test
INFO:__main__:Fitting model...
INFO:__main__:Finished fitting
INFO:__main__:Elasticnet model (alpha=0.500000, l1_ratio=0.500000):
INFO:__main__:  RMSE: 0.793164022927685
INFO:__main__:Logging parameters to MLflow
INFO:__main__:Logging trained model
MLflow run_id: 565a221aa54348929ca8da07edf1973a
INFO:botocore.credentials:Found credentials in environment variables.
Registered model 'ElasticnetWineModel' already exists. Creating a new version of this model...
2025/11/26 20:06:06 INFO mlflow.store.model_registry.abstract_store: Waiting up to 300 seconds for model version to finish creation. Model name: ElasticnetWineModel, version 9
The S3 URI of the logged model: s3://mlflow/4/565a221aa54348929ca8da07edf1973a/artifacts/model
Created version '9' of model 'ElasticnetWineModel'.
```

Expected output:

```text
INFO:__main__:Using MLflow tracking URI: http://mlflow-server.local
INFO:__main__:Using MLflow experiment: mlflow-minio-test
2023/08/16 13:45:22 INFO mlflow.tracking.fluent: Experiment with name 'mlflow-minio-test' does not exist. Creating a new experiment.
(More logs...)

MLflow run_id: fca4fffeab3d44e98ad2584f9f32a45
Successfully registered model 'ElasticnetWineModel'.
The S3 URI of the logged model: s3://mlflow/7/fca4fffeab3d44e98ad2584f9f32a45a/artifacts/model
```
Note that the S3 URI of the logged model vary. 

Navigate to the MLflow service UI at [http://localhost:5000/](http://localhost:5000/) (ported from the http://mlflow-server.local URL inside the cluster),
and you should see your run under the experiment "mlflow-minio-test". You can browse the run parameters, metrics and artifacts. For example: 

* Training hyperparameters and evaluation metrics:

<img src="images/mlflow-logging.png" width="1000"/>

You may notice that the "Metrics" and "Parameters" field are hidden by default, you can make them visible by clicking the "Columns" tab:

<img src="images/mlflow-show-columns.png" width=1000 />

When clicking the Run Name, we can also check where the model and other related files have been uploaded:

<img src="images/mlflow-uploaded-artifacts.png" width=1000 />

In this case, the model (which is a Pickle file) and its related files (such as the model dependency requirements) have been uploaded to the MinIO service. Navigate to [http://localhost:9001/](http://localhost:9001/) (ported from the http://mlflow-minio.local URL inside the cluster) and login using "minioadmin" as both the username and password, we can see there is a bucket named "mlflow":

<img src="images/minio-bucket-ui.png" width=1000 />

clicking the bucket (and its underlying folders) we can see the model and its related artifacts reside in the "mlflow" bucket:

<img src="images/minio-model-artifacts.png" width=1000 />

* Finally, we can also see the model has been registered to MLflow:

<img src="images/mlflow-model.png" width="1000"/>

# MLflow exercise
**Please do the exercise using the `mlops_eng` environment.**

In this exercise, you will train a LightGBM regression model to predict public bike sharing demand given attributes like datetime and weather conditions. The dataset ("bike_sharing_demand.csv", located under the same directory as this notebook) used in this exercise is a preprocessed version of [this kaggle dataset](https://www.kaggle.com/competitions/bike-sharing-demand/overview). You'll learn more about data preprocessing next week. Additionally, you'll use MLflow to track the model training and Deepchecks to evaluate the trained model. 

```python
import pandas as pd
from lightgbm import LGBMRegressor
import mlflow
from deepchecks.tabular import Suite
from deepchecks import SuiteResult, CheckResult
from deepchecks.tabular.checks import TrainTestPerformance, ModelInferenceTime, MultiModelPerformanceReport
from deepchecks.tabular import Dataset
from pathlib import Path
import os
import pickle
import boto3
import logging
import warnings

from typing import Tuple, Dict, Any

PARENT_DIR = Path("").parent.absolute()

# Suppress boto3 logging
boto3.set_stream_logger(name='botocore.credentials', level=logging.ERROR)
```

```python
def delete_file_if_existing(filename: str):
    """
    Delete a file if it exists
    """
    if os.path.exists(filename):
        print(f"Delete the existing {filename}")
        os.remove(filename)

```

## Download and split the dataset
Please note that the dataset (bike_sharing_demand.csv) contains data collected from two years so you'll find duplicated data points for most of the timestamps (hour-day-month) if you explore the dataset. 

### Load the data
First, implement the `pull_data` function that loads a CSV as a Pandas DataFrame from a given location.

<!-- grade_id: cell-416028b27a69c2cb, points: 0, grade: false, solution: true -->
```python
def pull_data(dataset_path: Path) -> pd.DataFrame:
    """
    Download the data set from a given path
    Args: 
        dataset_path (Path): Path of the CSV 
    Returns:
        A Pandas DataFrame of the dataset
    """
    ### START CODE HERE
    data = pd.read_csv(dataset_path)
    return data
    ### END CODE HERE

```

<!-- grade_id: cell-85e9a7af94c49319, points: 1, grade: true, solution: false -->
```python
# You can use this code cell to check if your pull_data function works correctly
dataset_path = PARENT_DIR / "bike_sharing_demand.csv"
df = pull_data(dataset_path)
assert df.shape == (10886, 12)

```

```python
# Show a concise summary of the DataFrame
df.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 10886 entries, 0 to 10885
Data columns (total 12 columns):
 #   Column      Non-Null Count  Dtype  
---  ------      --------------  -----  
 0   season      10886 non-null  int64  
 1   holiday     10886 non-null  int64  
 2   workingday  10886 non-null  int64  
 3   weather     10886 non-null  int64  
 4   temp        10886 non-null  float64
 5   atemp       10886 non-null  float64
 6   humidity    10886 non-null  int64  
 7   windspeed   10886 non-null  float64
 8   count       10886 non-null  int64  
 9   hour        10886 non-null  int64  
 10  day         10886 non-null  int64  
 11  month       10886 non-null  int64  
dtypes: float64(3), int64(9)
memory usage: 1020.7 KB
```

<details>
    <summary>Expected output</summary>
    <img src="./images/dataset-info.png"/>
</details>

Below is the explanation of each column in the dataset:

**Variables**:

| Column name |  Explanation | type |
|-------------|---------------|----|
| season      | 1 = spring, 2 = summer, 3 = fall, 4 = winter | integer
| holiday     | whether the day is considered a holiday | integer
| workingday  | 1 if day is neither weekend nor holiday, otherwise 0. | integer
| weather     | 1: Clear, Few clouds, Partly cloudy, Partly cloudy; 2: Mist + Cloudy, Mist + Broken clouds, Mist + Few clouds, Mist; 3: Light Snow, Light Rain + Thunderstorm + Scattered clouds, Light Rain + Scattered clouds; 4: Heavy Rain + Ice Pallets + Thunderstorm + Mist, Snow + Fog | integer
| temp        | temperature in Celsius | float
| atemp       | "feels like" temperature in Celsius | float
| humidity    | relative humidity | integer
| windspeed   | wind speed | float
| hour        | the hours of the datetime| integer
| day         | the day of the datetime| integer
| month       | the month of the datetime| integer

**Targets**: 

| Column name | Explanation                                     | Type
|-------------|-------------------------------------------------| ----|
| count       | number of total rentals                         | integer

### Split the data into train and test DataFrames
Then implement the `splitData` function that splits the dataset into a training and a test dataset, using the last 168 rows of the dataset as the test data.

<!-- grade_id: cell-8d442ea782a882c5, points: 0, grade: false, solution: true -->
```python
def splitData(input_df: pd.DataFrame) -> Tuple[pd.DataFrame, pd.DataFrame]:
    """
    Split a DataFrame into training and testing sets
    Args:
        input_df (DataFrame): The DataFrame to be split
    Returns:
        A tuple of training and testing DataFrame
    """
    ### START CODE HERE
    train_df = input_df[:-168]
    test_df = input_df[-168:]
    return train_df, test_df
    ### END CODE HERE
```

```python
df = pull_data(dataset_path)
train, test = splitData(df)
```

<!-- grade_id: cell-7ef629a52b6d83e4, points: 1, grade: true, solution: false -->
```python
# Check if train and test DataFrames are split correctly
expected_train_shape = (10718, 12)
expected_test_shape = (168, 12)

assert (
    train.shape == expected_train_shape
), "The dimension of the training dataset is not correct"
assert (
    test.shape == expected_test_shape
), "The dimension of the testing dataset is not correct"

expected_columns = [
    "season",
    "holiday",
    "workingday",
    "weather",
    "temp",
    "atemp",
    "humidity",
    "windspeed",
    "count",
    "hour",
    "day",
    "month",
]
assert set(train.columns) == set(
    expected_columns
), "The columns of the training dataset are not correct"
assert set(test.columns) == set(
    expected_columns
), "The columns of the training dataset are not correct"
```

```python
# Split the training and testing DataFrames into features and targets
target = "count"
input_df = pull_data(PARENT_DIR / "bike_sharing_demand.csv")
train, test = splitData(input_df)
train_x = train.drop([target], axis=1)
test_x = test.drop([target], axis=1)
train_y = train[[target]]
test_y = test[[target]]
```

```python
# Export the training and testing DataFrames for later use in the next notebooks
train.to_csv(PARENT_DIR / "train.csv", index=False)
test.to_csv(PARENT_DIR / "test.csv", index=False)
```

```python

```